# MT2505 Computing Project Part 1 (Semester 2 2021/22)

## Preliminaries

This is the first of three Computing Projects for MT2505 Abstract Algebra. This project focuses on demonstrating properties of binary operations using Python.

#### Before starting this project
You should have attempted the "List and Functions" worksheet, and the MT2000 Computing Workshop prior to this.

#### Submitting your project

The deadline for this part of the project is **5pm on Friday 4th March**. There are 10 marks available in this part, and it is worth 5% of your final grade in the module.

In order to submit your project, you must download it from this server **as a notebook** using `File> Download As> Notebook (.ipynb)`. You can then upload it to the relevant area on MySaint.

Your notebook must produce the correct results when run from top to bottom with a fresh Python kernel. To test this, you should click `Kernel> Restart & Run All`. You should also **validate** your project before submitting it (see below).

The process for submitting this part of your project is therefore
1. Use `Kernel> Restart & Run All` to ensure that your code works correctly when run in sequence.
2. Click the `Validate` button in the toolbar to ensure that your project passes the included validation cells.
3. Download your project using `File> Download As> Notebook (.ipynb)`.
4. Upload this file to MySaint.

#### Validation

The project already contains a number of code cells marked `# VALIDATION CELL`. You do not need to, and should not, make any changes to these cells.  Running one of these cells will perform some simple checks on the format of your work (for example, checking whether you have defined the function a question asks for). If a problem is detected an error will be produced giving you some feedback; if there is no detected problem there will be no output. **These cells do not test if your answer is actually correct**. You can run all of the validation cells by clicking the `Validate` button on the menu bar, but be aware that the validation for questions you have not attempted yet will automatically fail.

#### Getting help
Help is available during your scheduled computing classes.

#### Guidance

There is a [guidance document](https://moody.st-andrews.ac.uk/moodle/pluginfile.php/1376225/mod_resource/content/1/build/zz-index.html) on Moodle which you should read before attempting the project. This contains important information which may affect your grade.

# Binary Operations

As seen in lectures, we can represent binary operations by their Cayley tables. The Cayley table

<center>
$
\begin{array}{l|lll}
\ast & 0 & 1 & 2 \\ \hline
0    & 0 & 1 & 0 \\ 
1    & 0 & 0 & 0 \\
2    & 0 & 1 & 1 \\
\end{array}$
</center>

represents a binary operation where, for example, $2 \ast 1 = 1$ and $1 \ast 2 = 0$. The entry in row $x$ and column $y$ is the value of $x \ast y$.

The binary operation given by the table above could be represented in Python by a list of lists like so:

```python
X = [[0, 1, 0], [0, 0, 0], [0, 1, 1]]
```

Note that each of the nested lists corresponds to a row of the Cayley table, and the product $x \ast y$ is given by `X[x][y]`. 

```python
X[0][1]
```

Output:
```
1
```

This is a special property which only holds because the Cayley table is defined on the set $\{0, 1, \ldots, n - 1\}$ for some positive integer $n$ (where in this case $n = 3$). In this project, we will only consider binary operations on such a set $\{0, 1, \ldots, n-1\}$.

#### Question 1.

For the purposes of this project, a binary operation is a list `X` such that:
1. The length of `X` is a positive integer $n$.
2. Each entry of `X` is a list of length $n$, called a *row*.
3. Each row of `X` contains only `int` values $x \in \{0, 1, \ldots, n - 1\}$.

*The function `validate_binary_operation` defined below checks some of these properties. Complete it so that it correctly returns `True` if `X` is a valid binary operation, and raises an appropriate `TypeError` or `ValueError` otherwise.*


*Note*: the validation cell only contains a small number of tests for your function; it would be sensible to test it on more examples.

<p style='text-align:right;'> <b> [1 Mark] </b> </p>

```python
def validate_binary_operation(X):
    # Binary operations must be lists:
    if not isinstance(X, list):
        raise TypeError("binary operations should be lists, but got a " + type(X).__name__)

    # Binary operations must have positive length:
    if not len(X) > 0:
        raise ValueError("binary operations should be positive-length, but got a list of length 0")

    for row in X:
        # Each entry of a binary operation must be a list:
        if not isinstance(row, list):
            raise TypeError("binary operations should be lists of lists, but got a list containing a "
                           + type(X).__name__)
    
    # Complete the checks in this function!
    for row in X:
        if len(row)!= len(X):
            raise ValueError("each row should be equal to the length of X")
        for a in row:
            if type(a) != int:
                raise TypeError("rows should countain only integers")
                
            if a<0 or a>(len(X)-1):
                raise ValueError("each row should contain only values [0,1,...n-1]")

    return True

```

<!-- grade_id: cell-bde475e7e85f00f5, points: 1, grade: true, solution: false -->
```python
# VALIDATION CELL
# Cells like this contain some validation.
# If an error is produced, you need to fix something in your code.
# If no error is produced, this is *NOT* a guarantee you have correctly answered the question.
# Do not modify these cells.
from nose.tools import *
from sampleops import *
if not "validate_binary_operation" in globals():
    raise NotImplementedError("validate_binary_operation has not been defined in Question 1")

assert_raises(TypeError, validate_binary_operation, "banana")
assert_raises(TypeError, validate_binary_operation, ["a"])
assert_raises(ValueError, validate_binary_operation, [[1]])
assert_equal(validate_binary_operation([[0, 1], [1, 1]]), True)

```

#### Question 2.

Recall that a binary operation $\ast$ on a set $A$ is *associative* if $a\ast(b \ast c) = (a \ast b)\ast c$ for all $a,b,c \in A$. If you are asked whether an operation is associative, you should either prove that it is, or provide a counterexample to show that it is not. A very direct way of proving it would be to check every choice of three values $a, b, c \in A$, but this quickly becomes too much work to do by hand. For example, if $|A| = 50$, you would need to check more than one hundred thousand triples - and each one requires you to compute 4 products! This sort of direct verification is best done by computer.

*Write a function `is_associative_operation` which takes a binary operation (as defined above) `X` and returns `True` if the operation is associative, and `False` if not. Your function should first call `validate_binary_operation(X)` so that an appropriate `Exception` is raised if the argument is not a binary operation.*


<p style='text-align:right;'> <b> [2 Marks] </b> </p>

*Note:* recall that the product of `a` and `b` under `X` is given by `X[a][b]`.

Running `from sampleops import *` will give you access to two lists `associative_examples` and `non_associative_examples` which might help you test your function. Each entry of these lists is a binary operation, so for example you could test `is_associative(associate_examples[0])`. The validation cell tests your function on these lists.


```python

def is_associative(X):
    validate_binary_operation(X)
    for i in range(len(X)):
        for a in range(len(X)):
            for b in range(len(X)):
                if X[i][X[a][b]] != X[X[i][a]][b]:
                    return False
    return True

```

<!-- grade_id: cell-87d66870b6c320f8, points: 2, grade: true, solution: false -->
```python
# VALIDATION CELL
from nose.tools import *
from sampleops import *
if not "is_associative" in globals():
    raise NotImplementedError("is_associative has not been defined in Question 2")
    
assert_raises(TypeError, is_associative, "banana")

for test_op in associative_examples:
    assert_true(is_associative(test_op))
for test_op in non_associative_examples:
    assert_false(is_associative(test_op))

```

#### Question 3.

Recall that in general, a binary operation $\ast$ on a set $A$ has *identity* $e$ if $e\ast a = a \ast e = a$ for all $a \in A$. In the context of this practical, $e \in \{0, 1, \ldots, n-1\}$ is an identity of a list-of-lists binary operation `X` if and only if the row and column with index $e$ are both equal to `[0, 1, ..., n - 1]`.


*Write a function `identity` which takes a binary operation `X` and returns the identity of the binary operation, or `-1` if there is no identity. As for `is_associative`, an appropriate `Exception` should be raised if the input is not a binary operation.*

<p style='text-align:right;'> <b> [2 Marks] </b> </p>

*Note:* the row indexed $i$ can be obtained by `X[i]`, and the column indexed $i$ by `[row[i] for row in X]`. The list `[0, 1, ..., n-1]` can be obtained as `list(range(n))`.

Running `from sampleops import *` will give you access to two lists `identity_examples` and `identity_results` which might help you test your function. These lists have the property that `identity(identity_examples[i])` should be equal to `identity_results[i]`. The validation cell tests this property.

```python

def identity(X):
    counter = 0 
    validate_binary_operation(X)
    for i in range(len(X)):
        for a in range(len(X)):
            if X[a][i]==X[i][a] and X[a][i] == a:
                counter +=1
                if counter == len(X):
                    return i
        counter = 0
    return -1

```

<!-- grade_id: cell-aa92dc6ece4219c6, points: 2, grade: true, solution: false -->
```python
# VALIDATION CELL
from nose.tools import *
from sampleops import *
if not "identity" in globals():
    raise NotImplementedError("identity has not been defined in Question 3")
    
assert_raises(TypeError, identity, "banana")

for (test_op, test_res) in zip(identity_examples, identity_results):
    assert_equal(identity(test_op), test_res)

```

#### Question 4.

In order to find whether a binary operation defines a group on a set, we need to determine whether every element has an inverse under the binary operation.

*Write a function `has_inverses` which takes a binary operation `X` on $\{0, 1, \ldots, n-1\}$ and returns whether every $x \in \{0, \ldots, n-1\}$ has an inverse under `X`. An appropriate `Exception` should be raised if the input is not a binary operation.*

<p style='text-align:right;'> <b> [2 Marks] </b> </p>

*Note:* if `identity(X)` returns `-1`, then `has_inverses(X)` should return `False`. You may find it easier to write a function `has_inverse(X, i)` first, to determine whether a particular element has an inverse.

Running `from sampleops import *` will give you access to two lists `inverse_examples` and `inverse_results` which might help you test your function. These lists have the property that `has_inverses(inverse_examples[i])` should be equal to `inverse_results[i]`. The validation cell tests this property.

```python

def has_inverse(X , i):
    ident = identity(X)
    for a in range(len(X)):
        if X[a][i]==X[i][a] and X[a][i]==ident:
            return True
        elif a == (len(X)- 1):
            return False
            

def has_inverses(X):
    if identity(X)==-1:
        return False
    validate_binary_operation(X)
    for i in range(len(X)):
        if has_inverse(X,i)==False:
            return False
    return True

```

<!-- grade_id: cell-2b9ff9a80f8f475c, points: 2, grade: true, solution: false -->
```python
# VALIDATION CELL
from nose.tools import *
from sampleops import *
if not "has_inverses" in globals():
    raise NotImplementedError("has_inverses has not been defined in Question 4")

assert_raises(TypeError, has_inverses, "banana")

for (test_op, test_res) in zip(inverse_examples, inverse_results):
    assert_equal(has_inverses(test_op), test_res)

```

#### Question 5.

Given natural numbers $a$ and $b$ with $b > 0$, we will write $a \bmod b$ for the remainder on dividing $a$ by $b$, i.e. the number $r \in \{0, 1, \ldots, b - 1\}$ such that $a = bq + r$ for some $q \in \mathbb{N}$. We also write $\left\lfloor \frac{a}{b} \right\rfloor$ for this corresponding quotient $q$.

Define the binary operation $\ast$ on the set $\{0, 1, \ldots, 71\}$ as follows: 
for any $x, y \in \{0, 1, \ldots, 71\}$, let 
\begin{align*}
x_1 &= \left\lfloor\frac{x}{6}\right\rfloor \\
y_1 &= \left\lfloor\frac{y}{6}\right\rfloor \\ 
x_2 &= x \bmod 6 \\ 
y_2 &= y \bmod 6.
\end{align*}
Then define
$$x \ast y = 6\left([x_1 + y_1] \bmod 12\right) + \left([x_2y_2 + x_2 + y_2] \bmod 7\right).$$

*Store this binary operation in a variable `q5_operation`. Does this define an abelian group?*

<p style='text-align:right;'> <b> [3 Marks] </b> </p>

*Note:* your final answer should be clearly justified and written in a text cell; use `Cell>Cell Type>Markdown` to convert a code cell to a text cell. You may provide any mix of mathematical and computational justification.

You can get $\left\lfloor \frac{a}{b} \right\rfloor$ in Python using integer division `a // b`, and $a \bmod b$ using `a % b`. 

It would be sensible to test that your `q5_operation` does define a binary operation using `validate_binary_operation`.

```python
set_list = list(range(72))
q5_operation = []
for i in range(72):
    q5_operation.append([])
    for a in range(72):
        x1 = i//6
        y1 = a//6
        x2 = i%6
        y2 = a%6
        q5_operation[i].append(6*(x1+y1)%12+((x2*y2 +x2+y2)%7))
validate_binary_operation(q5_operation)

#Abelian group check
def abelian_check(X):
    for i in range(72):
        for a in range(72):
            if q5_operation[i][a]!=q5_operation[a][i]:
                return False
    return True

abelian_check(q5_operation)
```

Output:
```
True
```

<!-- grade_id: cell-fd9fbf75edf90452, points: 1, grade: true, solution: false -->
```python
# VALIDATION CELL

from nose.tools import *
from sampleops import *
if not "q5_operation" in globals():
    raise NotImplementedError("q5_operation has not been defined in Question 5")

```

The binary operation * on the set [0,1,...,71] defines an abelian group. This is shown by the computed check under the question code.
